---
title: "Tarefas: Controle de fluxo e repetições"
description: "Exercícios do curso 'Introdução à programação com R'"
encoding: "UTF-8"
output: 
  learnr::tutorial:
    css: css/custom.css
    language: pt
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
params:
  url_form: "https://docs.google.com/forms/d/e/1FAIpQLSfWCbu_4-jD7H-wV3gTxa80Q0Ko5Aylso3JOBZzRUDhwIKm3A/viewform"
---

```{r setup, include=FALSE}
library(learnr)
library(gradethisBR)
Sys.setenv(LANG = "pt")

knitr::opts_chunk$set(echo = FALSE, exercise.reveal_solution = FALSE)


gradethisBR::gradethis_setup(
  code_correct =  "Correto! Parabéns!",
  pass = "Correto! Parabéns!"#,
  #fail =  "Incorreto. Tente novamente! "
)


```


**Exercícios**

> Boas vindas! 
>
> Este material faz parte do curso [Introdução à programação com R](https://curso-r.com/cursos/intro-programacao/), da [Curso-R](https://www.curso-r.com/).
>
> Caso tenha dúvidas sobre como resolver os exercícios,  pedimos que escreva no [<i class="fab fa-discourse"></i> Discourse](https://discourse.curso-r.com/).


## Loops

Neste exercício, usaremos o pacote beepr.

```{r echo=TRUE}
# install.packages('beepr')
library(beepr)
```



1. Quando programamos o despertador do nosso celular para acordar de manhã, estamos, internamente, programando um laço (for) para que uma mesma tarefa (despertar) seja executada várias vezes.

Vamos imaginar que todos os dias o nosso celular verifica se ele precisa despertar às 8 horas da manhã. Se quisermos acordar todos os dias nesse horário, isso poderia ser representado, no R, pelo código abaixo.

**Observação**: Se você estiver acessando este tutorial em https://curso-r-tarefas.shinyapps.io/intro-programacao_4-loops/ a execução do código abaixo pode não funcionar. Experimente executar o código no RStudio instalado na sua máquina.

```{r loops-exemplo, exercise = TRUE}
#install.packages("beepr")

dias_da_semana <- c("seg", "ter", "qua", "qui", "sex", "sab", "dom")

for (dia in dias_da_semana) {
  mensagem <- paste0("8 horas da manhã de ", dia, "! Hora de acordar!")
  print(mensagem)
  beepr::beep(1)
  Sys.sleep(1)
}

```



O comando `Sys.sleep(1)` faz o R esperar 1 segundo antes de continuar executando o código. (Esse 1 segundo representano nosso exemplo as 24 horas entre dois dias).

1) Complete o código abaixo para  fazer o nosso despertador não "tocar" no fim de semana.

```{r loops-1, exercise = TRUE}
dias_da_semana <- c("seg", "ter", "qua", "qui", "sex", "sab", "dom")

for (dia in dias_da_semana) {
  if (..........) {
    mensagem <- paste0("8 horas da manhã de ", dia, "! Hora de acordar!")
    print(mensagem)
    beepr::beep(1)
    Sys.sleep(1)
  }
}
```

```{r loops-1-solution}
dias_da_semana <- c("seg", "ter", "qua", "qui", "sex", "sab", "dom")

for (dia in dias_da_semana) {
  if (!dia %in% c("sab", "dom")) {
    mensagem <- paste0("8 horas da manhã de ", dia, "! Hora de acordar!")
    print(mensagem)
    beepr::beep(1)
    Sys.sleep(1)
  }
}
```

```{r loops-1-check}
grade_this({
  if (identical(.result, .solution)) {
    pass("Correto!")
  }
  fail()
})
```

2) Complete o código abaixo para fazer o nosso despertador iterar (caminhar) no  dia do mês e não mais no dia da semana. Pode considerar que todo mês tem 30 dias.


```{r loops-2, exercise = TRUE}

dias_do_mes <- ..........

for (dia in ..........) {
  mensagem <- paste0("8 horas da manhã do dia ", ...)
  print(mensagem)
  beepr::beep(1)
  Sys.sleep(0.5)
}
```

```{r loops-2-solution}
dias_do_mes <- 1:30

for (dia in dias_do_mes) {
  mensagem <- paste0("8 horas da manhã do dia ", dia)
  print(mensagem)
  beepr::beep(1)
  Sys.sleep(0.5)
}
```

```{r loops-2-check}
grade_this({
  if (identical(.result, .solution)) {
    pass("Correto!")
  }
  fail()
})
```


## Enviar as respostas

Se você completou a tarefa e quer submeter, clique no botão abaixo ('Generate') para gerar o seu código hash e envie usando o formulário a seguir:

```{r context="server"}
learnrhash::encoder_logic()
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(ui_before = " ", ui_after = " ")
```

### Formulário para enviar o código

Caso o formulário abaixo não carregue, <a href='`r params$url_form`'>clique neste link</a> para abrir o formulário.


<iframe src="`r params$url_form`" width="640" height="673" frameborder="0" marginheight="0" marginwidth="0">Carregando…</iframe>



